<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Часть-1:-Много-гипотез" data-toc-modified-id="Часть-1:-Много-гипотез-1">Часть 1: Много гипотез</a></span><ul class="toc-item"><li><span><a href="#Мудрость" data-toc-modified-id="Мудрость-1.1">Мудрость</a></span></li><li><span><a href="#Тайлер" data-toc-modified-id="Тайлер-1.2">Тайлер</a></span></li><li><span><a href="#Рейтинги-президентов" data-toc-modified-id="Рейтинги-президентов-1.3">Рейтинги президентов</a></span></li><li><span><a href="#Сон" data-toc-modified-id="Сон-1.4">Сон</a></span></li><li><span><a href="#Say-my-name" data-toc-modified-id="Say-my-name-1.5">Say my name</a></span></li><li><span><a href="#Матан" data-toc-modified-id="Матан-1.6">Матан</a></span></li><li><span><a href="#Бутерброд" data-toc-modified-id="Бутерброд-1.7">Бутерброд</a></span></li></ul></li><li><span><a href="#Часть-2:-Продолжаем-проверять-гипотезы" data-toc-modified-id="Часть-2:-Продолжаем-проверять-гипотезы-2">Часть 2: Продолжаем проверять гипотезы</a></span><ul class="toc-item"><li><span><a href="#Однородная-игра" data-toc-modified-id="Однородная-игра-2.1">Однородная игра</a></span></li><li><span><a href="#Мыши-и-зеркало" data-toc-modified-id="Мыши-и-зеркало-2.2">Мыши и зеркало</a></span></li><li><span><a href="#Энергетик" data-toc-modified-id="Энергетик-2.3">Энергетик</a></span></li><li><span><a href="#Бомбардировки" data-toc-modified-id="Бомбардировки-2.4">Бомбардировки</a></span></li></ul></li><li><span><a href="#Еще-больше-гипотез" data-toc-modified-id="Еще-больше-гипотез-3">Еще больше гипотез</a></span><ul class="toc-item"><li><span><a href="#Анорексия" data-toc-modified-id="Анорексия-3.1">Анорексия</a></span></li><li><span><a href="#Метилфенидат" data-toc-modified-id="Метилфенидат-3.2">Метилфенидат</a></span></li><li><span><a href="#Сиэтл" data-toc-modified-id="Сиэтл-3.3">Сиэтл</a></span></li><li><span><a href="#Расход-топлива" data-toc-modified-id="Расход-топлива-3.4">Расход топлива</a></span><ul class="toc-item"><li><span><a href="#vs,-mpg" data-toc-modified-id="vs,-mpg-3.4.1">vs, mpg</a></span></li><li><span><a href="#am,-mpg" data-toc-modified-id="am,-mpg-3.4.2">am, mpg</a></span></li><li><span><a href="#vs,-am" data-toc-modified-id="vs,-am-3.4.3">vs, am</a></span></li></ul></li></ul></li></ul></div>

# Часть 1: Много гипотез

В этом домашнем задании нужно проверить много-много гипотез

Каждая дает по 5 баллов

**Часть 1**

| Задание              | Баллы |
|----------------------|-------|
| Мудрость             |     5 |
| Тайлер               |     5 |
| Рейтинги президентов |     5 |
| Сон                  |     5 |
| Say my name          |     5 |
| Матан                |     5 |
| Бутерброд            |     5 |
| Всего            |     35 |


**Часть 2**

| Задание            | Датасет         | Баллы |
|--------------------|-----------------|-------|
| Однородная игра    | killer_time     |     5 |
| Мыши и зеркало     | mouses_mirror   |     5 |
| Энергетик          | energy drink    |     5 |
| Бомбардировки      | london          |     5 |
| Анорексия          | anorexia        |     5 |
| Метилфенидат       | methylphenidate |     5 |
| Сиэтл              | seattle         |     5 |
| Расход топлива - 1 | mtcars: vs mpg  |     5 |
| Расход топлива - 2 | mtcars: am mpg  |     5 |
| Расход топлива - 3 | mtcars: am vs   |     5 |
|                    | Всего           |    50 |


```python
import numpy as np
import pandas as pd

from scipy import stats
from scipy import stats as sts

import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('ggplot')

# это на всякий случай для маков, чтобы скачался датасет
import ssl
ssl._create_default_https_context = ssl._create_unverified_context

# тут лежат все датасеты
ROOT = 'https://raw.githubusercontent.com/ilyaaaaaaaa/datasets_for_ds/main/data/'
# ROOT = '../data/'

```

## Мудрость

> Помните, что для человека звук его имени является самым сладким и самым важным звуком человеческой речи. (Дейл Карнеги)

Доверяй, но проверяй. Решим задачку по психологии! Все цитаты в пабликах вконтакте делятся на два вида: мудрые и недостаточно мудрые. Если цитата мудрая, то её можно лайкнуть. Поговаривают, что если в цитате есть личные местоимения, то пользователь автоматически сопоставляет себя автору этой цитаты и вероятность того, что она будет лайкнута, повышается. Например, цитату 

> Можно простить человеку все, кроме отсутствия. (Экзюпери) 

мне не очень хотелось бы лайкать. А вот цитата 

> Мне наплевать, что вы обо мне думаете. Я о вас вообще не думаю. (Булгаков) 

совершенно другое дело. 

В табличке `df_vk_wise` лежат данные о мудрых цитатах из [мудрого паблика.](https://vk.com/ooioioioioioo) В колонке `I` находится переменная, которая принимает значение `True`, когда в тексте цитаты было личное местоимение "Я".

Срежем первые $50$ постов, так как под ними количество репостов и лайков ещё окончательно не выставлено, и будем работать с остатками. 

```python
df = pd.read_csv(ROOT+'df_vk_wise.csv')[50:]
df.head()
```

Постройте гистограмму для числа лайков постов, содержащих личные местоимения и не содержащих. Как думаете, есть ли с распределением лайков какие-то проблемы?

<!-- grade_id: cell-5b1acf05ccd983b3, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

I_true = df.loc[df['I'] == True, 'likes']
I_false = df.loc[df['I'] == False, 'likes']

sns.histplot(I_true, bins=15, stat='density', color = 'lightskyblue')
sns.histplot(I_false,  bins=15, stat='density', color = 'lightcoral')
```

Да, проблемы есть. Это выбросы. Хвосты на гистограммах оказываются очень длинными. Найдите медианы и средние для числа лайков. Сравните их между собой. 

<!-- grade_id: cell-84afb85a40373b11, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

#медианы
I_true_median = I_true.median()
I_false_median = I_false.median()

#средние
I_true_mean = I_true.mean()
I_false_mean = I_false.mean()

print(I_true_median, I_false_median,I_true_mean, I_false_mean)

#В среднем число лайков для цитат с Я немного больше, как и медиана
```

Проверьте гипотезу о равенстве среднего числа лайков. Используйте для этого $z$-тест, основанный на ЦПТ. Дисперсию расчитывайте по выборке, используя несмещённую оценку `(ddof=1)`. Запишите наблюдаемое значение статистики в переменную `z_obs`. 

<!-- grade_id: cell-1acfb89f07181a0a, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

alpha = 0.05 

n_I_true = I_true.size
n_I_false = I_false.size
std = np.sqrt(I_true.var(ddof=1)/n_I_true + I_false.var(ddof=1)/n_I_false)
z_obs = (I_true_mean - I_false_mean) / std
```

Найдите критическое значение статистики. Сравните его с наблюдаемым. Рассчитайте p-value для ситуации с двусторонней альтернативой. Сделайте вывод о выдвинутой гипотезе. Критическое значение статистики запишите в переменную `z_crit`, а p-значение в переменную `pval`. 

<!-- grade_id: cell-9ebff976022e649a, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
z_crit = stats.norm.ppf(1 - alpha)
pval = 2*(1-stats.norm.cdf(np.abs(z_obs)))

print(f"Наблюдаемое значение: {z_obs:.7}")
print(f"Критическое значение: {z_crit:.7}")
print(f"P-значение: {pval:.3}") 

```

```python
# проверка того, что задание решено верно
assert np.abs(z_obs - 2.6644610922505323) < 1e-5
```

Можно заметить, что у распределения числа репостов длинные хвосты. В данных есть выбросы, которые завышают среднее значение. Если выбросы сильные, то ЦПТ может ломаться. Из-за этого имеет смысл в данном случае проверять гипотезу о равенстве медиан, которые устойчивы к выбросам.

-------------------

## Тайлер

__Люди всё время меня спрашивают, знаю ли я Тайлера Дёрдена?__ Да, знаю. Он делает лучшее мыло в стране. 

__а)__ Тайлер считает партию изделий хорошей, если дисперсия размера не превышает $25$ грамм$^2$. Пока Тайлер летел в самолёте, он измерил $10$ кусков мыла. Выборочная дисперсия оказалась равна $27$ грамм$^2$. Можно ли продавать эту партию или её надо списать в брак? 

При проверке этой гипотезы найдите $p$-значение и запишите его в переменную `pval1`. Наблюдаемое значение статистики запишите в переменную `chi_obs`. Используйте уровень значимости $1\%$.

<!-- grade_id: cell-2e30e820f90148d5, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

alpha = 0.01
chi_crit = stats.chi2(n-1).ppf(1-alpha)

sigma2_null = 25
sigma2 = 27
n = 10
chi_obs = (n-1)*sigma2/sigma2_null 

pval1 = 1-stats.chi(n-1).cdf(chi_obs)

print(f'Наблюдаемое значение {chi_obs:.5}')
print(f"Критическое значение: {chi_crit:.7}")
print(f"P-значение: {pval1:.3}")

#pval1 меньше уровня значимости - гипотеза о том, что можно продавать, отвергается
```

__б)__ Для того, чтобы уменьшить количество бракованного мыла, Тайлер собирается внедрить новую технологию производства. Тайлер произвёл по старой технологии $100$ кусков мыла. Стандартное отклонение этой партии составило $6$ грамм. Для  новой технологии стандартное отклонение, посчитанное по $100$ кускам мыла, составило $5$ грамм. Правда ли, что новая технология обладает меньшей дисперсией и её нужно внедрить на производстве?

При проверке этой гипотезы найдите $p$-значение и запишите его в переменную `pval2`. Используйте уровень значимости $1\%$.

<!-- grade_id: cell-88a96be4941d7e62, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

n = 100
alpha = 0.01
std_old = 6
std_new = 5

F_obs = std_old**2/std_new**2

f_rv = stats.f(n - 1, n- 1) 
F_crit = f_rv.ppf(1 - alpha/2)  

pval2 = 2*(1-f_rv.cdf(F_obs))
print(f'Наблюдаемое значение {F_obs:.5}')
print(f"Критическое значение: {F_crit:.7}")
print(f"P-значение: {pval2:.3}")
#pval2 больше уровня значимости - нулевая гипотеза не отвергается
```

```python
# проверка того, что задание решено верно
assert np.abs(chi_obs - 9.72) < 1e-1
assert np.abs(F_obs - 1.44) < 1e-1

```

## Рейтинги президентов

В таблице `presidents.csv` представлены рейтинги различных президентов США. Социальные опросы проводились раз в квартал. Опрашивалась случайная выборка граждан. У каждого спрашивалось одобряет ли он деятельность президента.

__Люди всё время меня спрашивают, знаю ли я Никсона.__ Да, Ричард Никсон был президентом США с 20 января 1969 года по 9 августа 1974 года. 17 июня 1972 года произошёл Уотергейтский скандал. Давайте проанализируем, что при этом произошло с его рейтингом.

```python
df = pd.read_csv(ROOT+'presidents.csv')
df.head()
```

Чтобы привести таблицу к удобному виду используйте функцию `pd.melt`. Вытащите из таблички рейтинг Никсона с 1 квартала 1969 года по 2 квартал 1972 года (включая его). Запишите эти данные в переменную `x`. Рейтинг с 3 квартала 1972 года по 2 квартал 1974 года поместите в переменную `y`. Пропуски в данных удалите из выборки.

```python
df_melt = pd.melt(df, id_vars=['year'], value_vars=['Qtr1', 'Qtr2', 'Qtr3', 'Qtr4'],
                    var_name='Qtr', value_name='Rating')
```

<!-- grade_id: cell-0c35f4bdddbb0029, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

x = df_melt[df_melt['year'].between(1969, 1972)]['Rating'].dropna()
y = df_melt[df_melt['year'].between(1972, 1974)]['Rating'].dropna()

```

Правда ли, что после скандала средний рейтинг Никсона рухнул? Предполагайте, что рейтинг распределён нормально, а также, что дисперсия рейтинга не изменилась из-за скандала. Наблюдаемое значение статистики запишите в переменную `obs`, в переменную `pval` запишите $p$-значение.

<!-- grade_id: cell-468606ed00795265, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

#H0 - не изменился
#H1 - рухнул
alpha = 0.05 
s2 = ((x.size - 1)*x.var() + (y.size - 1)*y.var())/(x.size + y.size - 2)
diff_sd = np.sqrt(s2/x.size + s2/y.size)
obs = (x.mean() - y.mean())/diff_sd
t_crit = stats.t(df = x.size + y.size - 2).ppf(1 - alpha)

print(f'Наблюдаемое значение {obs:.3} больше критического {t_crit:.3} => нулевая гипотеза отвергается')
```

```python
from statsmodels.stats.weightstats import ttest_ind

t_obs, p_val, _ = ttest_ind(x, y, value=0, alternative='larger')

print(f"Наблюдаемое значение: {t_obs:.3}") 
print(f"P-значение: {p_val:.3}") 
#pval меньше уровня значимости -> H0 отвергается -> рейтинг рухнул
```

```python
# проверка того, что задание решено верно
assert np.abs(obs - 3.5640741313229793) < 1e-5

```

Снова президентом Никсон станет в $3000$ году. 

![](https://vignette.wikia.nocookie.net/futurama/images/2/2f/Richard_M._Nixon%27s_head.jpg/revision/latest/scale-to-width-down/250?cb=20131123190522&path-prefix=ru)

## Сон

Проводится эксперимент о пользе нового снотворного. Для его проведения используется группа из $10$ студентов. В первый день каждый из них ложится спать. На утро измеряется, сколько он проспал. Из результата вычитают $8$ (нормальное число часов сна). На второй день каждому дают таблетку. Утром снова измеряют сколько человек проспал, и вычитают $8$. Результаты эксперимента записаны в табличку `sleep.csv`.

```python
df = pd.read_csv(ROOT+'sleep.csv')
df.head()
```

Правда ли, что сон студентов до приёма снотворного не отличается от нормы (отклонение от 8 часов равно нулю). Проверьте эту гипотезу на $5\%$ уровне значимости, в переменную `pval1` запишите p-значение. 

<!-- grade_id: cell-c04ca7c5b2d46f59, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
alpha = 0.05
d = df[df['group'] == 1]['extra']
tobs1 = (d.mean()-8)/d.std(ddof=1)/np.sqrt(d.size)
pval1 = 2*(stats.t.cdf(-np.abs(tobs1), d.size-1))
```

Правда ли, что благодаря таблеткам студенты стали спать больше (односторонняя альтернатива)? Запишите в переменную `pval2` соотвествующее p-значение. 

<!-- grade_id: cell-4c7c2e8e8b3c5de8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
d1 = df[df['group'] == 2]['extra']
tobs2 = stats.ttest_rel(d, d1)[0]
```

<!-- grade_id: cell-2167298693b77a31, points: 1, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert np.abs(tobs1 - 1.325710) < 1e-5
assert np.abs(tobs2 - 4.062127) < 1e-5

```

## Say my name

__Люди всё время меня спрашивают, знает ли меня Барсик.__  Многие владельцы кошек уверены, что их питомцы могут понимать достаточно большое количество слов, а также узнавать собственное имя. Однако долгое время никто не пытался доказать это экспериментально. Разобраться с этой проблемой [решили учёные из Токийского университета.](https://nplus1.ru/news/2019/04/04/cat-name?utm_source=obrazovach&utm_medium=social&utm_campaign=yaponskie-uchenye-dokazali--chto-domashnie-k)

Было проведено четыре эксперимента с участием домашних кошек и кошек из кошачьих кафе. Учёные использовали метод привыкания-отвыкания (habituation-dishabituation), который часто используется в экспериментах с участием маленьких детей. 

В соответствии с гипотезой, которая лежит в основе этого метода, определить реакцию на какой-либо стимул можно, представив его после схожих с ним стимулов. К примеру, для того, чтобы выяснить, может ли младенец узнать свое имя, ему представляют несколько слов подряд: с каждым новым словом его реакция снижается, но появление имени должно (в случае, если ребенок действительно узнает свое имя) ее вернуть. Этот же метод применили и к кошкам: реакция кошек на слова характеризовалась поворотом головы, движением ушей, вокализацией, движениями хвоста или перемещениями.

В первом эксперименте исследователи проследили за реакцией кошек на собственное имя, произнесенное хозяином, после четырех существительных (16 кошек, живущих дома с двумя другими кошками или меньше). Данные об этом эксперименте находятся в табличке `cats_names.csv`. 

```python
df = pd.read_csv(ROOT+'cats_names.csv', sep='\t')
df.head()
```

Будем проверять гипотезу о том, что кошки понимают своё имя с помощью упрощённой процедуры: 

- Вытащим из таблички все строчки, где экспериментаторы называли $4-$ое существительное. Они помечены в колонке `stimuli` как `Noun 4`. Если в колонке `Total Score` стоит положительное число, будем считать, что кошка среагировала на слово. В вектор `noun4` запишем единицу. В противном случае — ноль. 
- Вытащим по аналогии информацию о том, откликнулась ли кошка на своё имя (пометка `Test Cat's Name`) в колонке `stimuli`. Запишем её в вектор `name`.
- Используем тест о равенстве пропорций для зависимых выборок, чтобы ответить на вопрос "А правда ли кошка откликается на своё имя?" В переменную `pval` запишем $p$-значение, полученное при проверке этой гипотезы. 

```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code
noun4 = [(1 if total_score > 0 else 0) for total_score in df[df['stimuli'] == 'Noun 4']['Total Score']]
name = [(1 if total_score > 0 else 0) for total_score in df[df['stimuli'] == "Test Cat's Name"]['Total Score']] 
pd.crosstab(noun4, name)
# здесь удобно создавать вот такую свободную таблицу
```

```python
n = 16
c = 1
b = 5
z_obs = (b-c)/np.sqrt(b+c - (b-c)**2/n)
pval = 2*(1-stats.norm.cdf(z_obs))
pval
```

<!-- grade_id: cell-4ab37a21e0f8a42c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-473051ae67ddc9a6, points: 1, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert sum(name) == 14
assert np.abs(z_obs - 1.7888543819998317) < 1e-5

```

Выводы учёных могут отличаться от наших. Результат проверки гипотезы зависит от многих внешних факторов. Например, от громкости произнесённого кошачьего имени. Все такие факторы надо как-то изолировать.

[В оригинальном исследовании](https://www.nature.com/articles/s41598-019-40616-4) учёные реализовывали более сложные статистические процедуры. Они пытались проконтролировать все внешние факторы, которые могли бы повлиять на результаты эксперимента. О том, как это можно делать, мы с вами поговорим в будущем. 

![](https://cdn.nplus1.ru/images/2019/04/04/47a222337ec0984bcec1644f12b4c328.gif)

## Матан

Лекции по матану на экономическом факультете читают два преподавателя: Злой и Добрый. Каждый год они чередуются. Среди студентов ходит слух, что в год, когда преподаёт Злой, отчисляют куда больше народу. В таблице `matan.csv` лежит информация по $418$ студентам экономического факультета. 

В столбце `is_angry` стоит единица, если этому студенту преподавал Злой. В столбце `target_1` стоит единица, если человек закончил первый курс. Проверьте гипотезу о том, что Злой отчисляет больше народу. В переменную `pval` запишите $p$-значение, получившееся при проверке гипотезы. 

```python
df = pd.read_csv(ROOT+'matan.csv', sep='\t')
df.head()
```

```python
df.dropna(inplace=True)
```

<!-- grade_id: cell-06da65b1aca523bb, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
angry = df[df['is_angry'] == 1]['target_1']
not_angry = df[df['is_angry'] == 0]['target_1']
z_obs = stats.ttest_ind(zloy, dobry, equal_var=False)[0]
#гипотеза о том, что злой отчисляет больше народу, не отвергается
```

```python
# проверка того, что задание решено верно
assert np.abs(np.abs(z_obs) - 0.29897)  < 1e-5

```

## Бутерброд

Вася решил проверить известное утверждение о том, что бутерброд падает маслом вниз. Для этого он провёл серию из $200$ испытаний. Бутерброд упал маслом вниз $105$ раз. Маслом вверх $95$ раз. Можно ли утверждать, что бутерброд падает маслом вниз так же часто, как и маслом вверх на уровне значимости $5\%$? В качестве ответа сохраните в `pval` $p$-значение. 

__Примечание:__ для поиска дисперсии используйте смещённый вариант. Либо по формуле $\frac{p\cdot(1-p)}{n}$ либо с помощью `std(ddof=0)`. При поиске `pval` обратите внимание, что алтернатива двусторонняя. А ещё [прочитайте историю](https://absentis.livejournal.com/27153.html) про то, почему Коперник мазал бутеры маслом. 

<!-- grade_id: cell-db59ee9982bc3985, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
#H0: бутерброд падает маслом вниз и маслом вверх с равной вероятностью p = 0.5
#H1: вероятности не равны
n = 200
m = 105
alpha = 0.05
p = 105/200 
q = 1 - p 

alpha = 0.05
p = 105/200 
pval = 2 * (1 - stats.norm.cdf(abs(z_obs))) 
z_obs = (p - 0.5) / np.sqrt(0.5 * 0.5 / n) 

#pval больше уровня значимости, значит, нулевую гипотезу не отвергаем и 
#бутерброд падает маслом вверх и вниз с одинаковой вероятностью
```

```python
# проверка того, что задание решено верно
assert np.abs(z_obs - 0.707) < 1e-3
```

<img src="https://i.ibb.co/Ks57rDV/student-teacher-dist.png" height="300" width="600"> 

# Часть 2: Продолжаем проверять гипотезы


## Однородная игра

Киллер - это весёлая лагерная игра. Краткие правила игры:

- Игроки пишут свои имена на карточках и кладут их в шляпу
- Игроки рандомно вытаскивают карточки из шляпы 
- Полученное имя - имя жертвы
- Каждый игрок одновременно и убийца и жертва
- Чтобы убить другого игрока, надо оказаться с ним наедине и застрелить из пальца
- Жертва убитого становится новой жертвой игрока
- Побеждает тот, кто совершит наибольшее число убийств

На экономическом факультете было два раунда игры в киллера. Все убийства регистрировались в гугл-форме. В табличках `killer_time1.csv` и `killer_time2.csv` лежит время в часах, которое проходило между убийствами.

```python
killer1 = pd.read_csv(ROOT+"killer_time1.csv").hours_between_kill.values
killer2 = pd.read_csv(ROOT+"killer_time2.csv").hours_between_kill.values
```

```python
plt.hist(killer1, bins=20);
```

```python
plt.hist(killer2, bins=20);
```

Распределение выглядит, как экспоненциальное. Более того, по гистограммам возникает ощущение, что это одинаковое распределение. С помощью теста Колмогорова-Смирнова проверьте гипотезу о том, что обе выборки пришли из одного и того же распределения, $p$-значение, полученное при проверке запишите в переменную `pvalue1`. 

<!-- grade_id: cell-f8134e5d39cba5cd, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue1 = stats.kstest(killer1, killer2)[1]
```

С помощью теста Колмогорова проверьте гипотезу о том, что время для первой игры имеет экспоненциальное распределение с параметром $\lambda = 0.026$, $p$-значение, полученное при проверке запишите в переменную `pvalue2`. 

<!-- grade_id: cell-da5a87e744840bab, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

lambd = 0.026
pvalue2 = stats.kstest(killer1, 'expon', args = (0, 1/lambd))[1]
```

<!-- grade_id: cell-be55b7fce581170d, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue1 > 0.5

```

## Мыши и зеркало

$16$ лабораторных мышей были помещены в двухкомнатные клетки, в одной из комнат висело зеркало. С целью установить, есть ли у мышей какие-то предпочтения насчет зеркал, измерялась доля времени, которое каждая мышь проводила в каждой из своих двух клеток.

![](https://ars.els-cdn.com/content/image/1-s2.0-S0168159104000061-gr1.gif)

В файле `mouses_mirror.txt` лежит информация о том, какую долю времени мышка проводила в клетке с зеркалом. 

```python
x = pd.read_csv(ROOT+'mouses_mirror.txt', header = None)[0].values
plt.hist(x, density=True);
```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью критерия знаков. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия. Запишите в переменную `pvalue_1` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-6b6cc3cc3f9d2e2d, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
#Ho = доля времени, проводимая мышкой в клетке с зеркалом, равна 0.5 
#Ha = не равна 0.5

from statsmodels.stats.descriptivestats import sign_test
# применяется для проверки равенства медианы некоторому заданному значению (для одной выборки)
# или о равенстве нулю медианы разности (для двух связанных выборок)
# подходит для небольших выборок;
# не требует знания распределения выборки.

n = x.shape[0]
med_0 = 0.5

T, pvalue_1 = sign_test(x, mu0=med_0)
```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью теста Вилкоксона. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия.  Запишите в переменную `pvalue_2` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-689716dd8c46567c, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
# применяется для сравнения медиан двух выборок;
# подходит для выборок любого размера;
# не требует нормальности выборки.

pvalue_2 = stats.wilcoxon(x - 0.5)[1]
```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью асимптотического $z$-теста. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия.  Запишите в переменную `pvalue_3` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-4ce505503163ae74, points: 0, grade: false, solution: true -->
```python
from statsmodels.stats.weightstats import ztest,zconfint
# основан на статистике z
# подходит для проверки гипотез о долях в больших выборках (обычно более 30 наблюдений);
# предполагает нормальность выборки.
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue_3 = ztest(x, value=0.5)[1]
```

<!-- grade_id: cell-2f9dfc7c1834c8a5, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue_1 < 0.05
assert pvalue_2 > 0.05
assert pvalue_3 > 0.05

```

Объясните, почему `pvalue` от теста к тесту ведёт себя именно так. Подробнее про опыт с мышками читайте [в оригинальном исследовании.](https://www.sciencedirect.com/science/article/abs/pii/S0168159104000061)

## Энергетик

Исследовательница Глафира считает, что любовь к энергетическим напиткам и успешность сдачи экзамена по математической статистике должны быть как-то связаны. Опросив $200$ своих однокурсников, она получила следующие результаты:

- Сдал, пьёт энергетик: 20
- Сдал, не пьёт энергетик: 120
- Завалил, пьёт энергетик: 40
- Завалил, не пьёт энергетик: 20

Правда ли, что энергетик оказывает на успешность сдачи экзамена какое-то влияние? Проверьте гипотезу с помощью критерия Пирсона. В качестве ответа запишите в переменную `pvalue` получившееся $p$-значение.

<!-- grade_id: cell-8806588bc0865379, points: 0, grade: false, solution: true -->
```python

```

```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
passed = np.array([20, 120])
fail = np.array([40, 20])
table = pd.DataFrame({'pass': passed, 'fail': fail}, 
                     index=['not_drink', 'drink'])
table
```

```python
chi, pvalue, df_, expected = stats.chi2_contingency(table.values, correction=False)
pvalue
```

<!-- grade_id: cell-920f09f17b0aafde, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue < 0.05

```

## Бомбардировки

В таблице `london_07-09-1940.tsv` лежат данные о бомбардировке Лондона 7 сентября 1940 года немецкой авиацией. Интересно выяснить, насколько прицельно велась бомбардировка. Если немецкие войска знали, где именно находятся британские военные объекты, они вели бы бомбардировку пристрельно.  

Ещё больше данных о бомбардировке Лондона ищите здесь: http://bombsight.org/data/sources/

```python
df = pd.read_csv(ROOT+'london_07-09-1940.tsv', sep='\t')
print(df.shape)
df.head()
```

```python
plt.scatter(df.x, df.y);
```

В колонках $x$ и $y$ лежат координаты точки, куда упала бомба. Разбейте Лондон решёткой размера $30 \times 30$ на $900$ небольших квадратных участков. С помощью функции `binned_statistic_2d` из пакета `scipy.stats`, подсчитайте как часто снаряд попадал на какой из участков. Поместите информацию об этом в переменную `ret`. 

<!-- grade_id: cell-2a06adb25a0b5542, points: 0, grade: false, solution: true -->
```python
binx = np.linspace(df.x.min(), df.x.max(), 31)
biny = np.linspace(df.y.min(), df.y.max(), 31)

### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

ret = stats.binned_statistic_2d(df.x, df.y, None, 'count', bins = [binx, biny])

sns.heatmap(ret.statistic);
```

Распакуйте информацию о попаданиях в участки в вектор `z`. Нарисуте гистограмму для получившегося распределения. 

<!-- grade_id: cell-1b692a90383e2eff, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

z = ret.statistic.flatten()
plt.hist(z, bins=range(10), alpha=0.5)
plt.xticks(range(10))
plt.xlabel('Число попаданий на участок')
plt.ylabel('Частота')
```

Участков много, а вероятность попадания на отдельный участок мала. При справедливости гипотезы о низкой точности стрельбы мы можем воспользоваться законом редких событий:

> Если в схеме испытаний Бернулли $p \to 0$, а $n \to \infty$, оно сходится к распределению Пуассона. 

Получается, что __при верности гипотезы о низкой точности стрельбы,__ число попаданий в квадрат имеет распределение Пуассона. Дополнительно будем предполагать, что попадания на разные участки независимы друг от друга. 

C помощью критерия Пирсона проверьте гипотезу о том, что число попаданий на участок имеет распределение Пуассона. Прежде, чем вычислять статистику, объедините все наблюдения из хвоста в одно по аналогии с лекцией. Хвостом будем считать все наблюдения $\ge 9$. Получившуюся в результате расчётов статистику запишите в переменную `T`.

<!-- grade_id: cell-ae799b5490316110, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

n = len(df)
p = z.sum() / (30**2 * n)
mu = n * p
expected = stats.poisson(mu).pmf(np.arange(10)) * len(z)
expected[8] = expected[8] + expected[9]
expected = np.delete(expected, 9)
```

```python
observed = np.histogram(z, bins=range(11))[0]
observed[8] = observed[8] + observed[9] 
observed = np.delete(observed, 9) 

T = np.sum((observed - expected)**2 / expected)
df = len(expected) - 1 # число степеней свободы
alpha = 0.05
critical_value = stats.chi2.ppf(q=1-alpha, df=df) 
#T > critical_value -> гипотеза о распределении Пуассона отвергается
```

<!-- grade_id: cell-2986a877274d9a7e, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert T > 6000

```

Какой вывод мы можем из этого сделать? Правда ли, что бомбардировка осуществлялась прицельно? 

Видимо, правда

# Еще больше гипотез

- Ниже приведены несколько датасетов, их необходимо самостоятельно проанализировать и ответить на определенные вопросы
- Для ответа на эти вопросы нужно сделать следующее
    - Сделать хорошую визуализацию исследуемых переменных
    - Для каждой из задач подобрать **все подходящие** параметрические и непараметрические критерии, проанализировать результаты этих критериев (тут имеет смысл обернуть все критерии в фунцию или словарь - как удобно)
- Для каждого набора данных проверьте нормальность выборки и допустимость использования параметрических критериев (можно прокомментировать гистограмму и воспользоваться соответствующими критериями согласия)
- Не забудьте оставить краткие комментарии к вашим действиям и сформулировать выводы, которые вы сделали в ходе анализа


## Анорексия

В файле anorexia.txt записан вес пациентов до начала терапии анорексии и после ее окончания. Была ли терапия в эффективна?


```python
df = pd.read_csv(ROOT+'anorexia.txt', sep='\t')
df.head()

```

```python
sns.histplot(df, bins=10, stat='density')
```

```python
stats.ks_2samp(df["Before"], df["After"])
#гипотеза о нормальности распределений не отвергается, можно воспользоваться парам.критериями
```

```python
#Ho - Терапия неэффективна
#Ha - Терапия эффективна
```

```python
from statsmodels.stats.descriptivestats import sign_test

def crit(df_B, df_A):
    alpha = 0.05
    pvalue_sign = sign_test(df_B, df_A)[1] #критерий знаков
    pvalue_ttest = stats.ttest_rel(df_B, df_A)[1]#t-критерий Стьюдентa для связанных выборок
    pvalue_t = stats.pearsonr(df_B, df_A)[1] #t-test
    wilcoxon_stat, pvalue_wilc = stats.wilcoxon(df_B, df_A) #Критерий Уилкоксона
    return pvalue_sign, pvalue_ttest, pvalue_t, pvalue_wilc

pval = crit(df["Before"], df["After"])
#гипотеза о неэффективности отвергается
```

## Метилфенидат
В рамках исследования эффективности препарата метилфенидат пациенты с синдромом дефицита внимания и гиперактивности в течение недели принимали либо метилфенидат, либо плацебо. В конце недели каждый пациент проходили тест на способность к подавлению импульсивных поведенческих реакций. На втором этапе плацебо и препарат менялись, и после недельного курса каждый испытуемые проходили второй тест. Был ли эффект от применения препарата? Данные находятся в файле methylphenidate.txt.



```python
df = pd.read_csv(ROOT+'methylphenidate.txt', sep=' ')
df.head()
```

```python
sns.histplot(df, stat='density', bins = 15)
```

```python
stats.ks_2samp(df['Placebo'], df['Methylphenidate'])
#гипотеза о нормальности распределений не отвергается, можно воспользоваться парам.критериями
```

```python
#Ho - Эффекта не было
#Ha - был эффект от применения препарата

from statsmodels.stats.descriptivestats import sign_test

def crit(df_B, df_A):
    alpha = 0.05
    stat, pvalue_ttest = stats.ttest_rel(df_B, df_A)#t-критерий Стьюдентa для связанных выборок
    pvalue_sign = sign_test(df_B, df_A)[1] #критерий знаков
    pvalue_t = stats.pearsonr(df_B, df_A)[1] #t-test
    wilcoxon_stat, pvalue_wilc = stats.wilcoxon(df_B, df_A) #Критерий Уилкоксона
    return pvalue_ttest, pvalue_sign, pvalue_t, pvalue_wilc


pval = crit(df['Placebo'], df['Methylphenidate'])
#гипотеза о неэффективности отвергается
```

## Сиэтл
В файле seattle.txt записаны цены на недвижимость в одном из районов Сиэтла в 2001 году и в 2002 году (объекты недвижимости выбирались случайно). Изменились ли в среднем цены в этом районе за год?


```python
df = pd.read_csv(ROOT+'seattle.txt', sep='\t')
df.head()
```

```python
fig, (ax1, ax2) = plt.subplots(nrows=1, ncols=2, figsize=(10, 5))
ax1.hist(df[df['Year'] == 2001]['Price'], bins=20)
ax1.set_title('Цены на недвижимость в 2001 году')
ax1.set_xlabel('Цена')
ax1.set_ylabel('Количество недвижимости')
ax2.hist(df[df['Year'] == 2002]['Price'], bins=20)
ax2.set_title('Цены на недвижимость в 2002 году')
ax2.set_xlabel('Цена')
ax2.set_ylabel('Количество недвижимости')
plt.show()
```

```python
stats.ks_2samp(df[df['Year'] == 2001]['Price'], df[df['Year'] == 2002]['Price'])
#гипотеза о нормальности распределений не отвергается, можно воспользоваться парам.критериями
```

```python
# H0: в среднем цены не изменились
# H1: изменились
```

```python
#z-тест
alpha = 0.05                                                    
x = df[df['Year'] == 2001]['Price']
y = df[df['Year'] == 2002]['Price']
nx, ny = x.size, y.size                                   

diff = x.mean() - y.mean()                                
diff_std = np.sqrt(x.var(ddof=1)/nx + y.var(ddof=1)/ny)   

z_obs = diff/diff_std                                           
z_crit = stats.norm.ppf(1 - alpha)                       

print(f'Наблюдаемое значение {z_obs:.5} меньше критического {z_crit:.3} => нулевая гипотеза не отвергается')
```

```python
def crit(df_P, df_M):
    alpha = 0.05
    stat, pvalue_tcrit = stats.ttest_ind(df_P, df_M) #t-критерий Стьюдентa
    mannwhitneyu_stat, pvalue_mann = stats.mannwhitneyu(df_P, df_M) #Критерий Манна-Уитни
    return pvalue_tcrit, pvalue_mann

pval = crit(df[df['Year'] == 2001]['Price'], df[df['Year'] == 2002]['Price'])
#нулевая гипотеза не отвергается
```

## Расход топлива
В файле mtcars.csv находятся данные из американского журнала Motor Trend 1974 года. Они описывают расход топлива в зависимости от 10 характеристик автомобиля (все автомобили 1973-1974 года). А именно, в данных вы найдете следующие столбцы:
1. mpg — расход топлива (миль/галлон); 
2. cyl — количество цилиндров мотора; 3) disp — объем мотор (куб. дюйм);
4. hp — мощность мотора (в лошадиных силах);
5. drat — передаточное число редуктора заднего моста;
6. wt — вес (тысяч фунтов);
7. qsec — время, за которое автомобиль проезжает 1/4 мили;
8. vs — тип двигателя (0 = V-образный, 1 = рядный);
9. am — тип коробки передач (0 = автоматическая, 1 = ручная);
10. gear — количество передних передач; 11) carb — количество карбюраторов.



- Влияет ли тип двигателя на расход топлива? 
- А тип коробки передач?
- Взаимосвязаны ли тип двигателя и тип коробки передач?

```python
df = pd.read_csv(ROOT+'mtcars.csv', index_col=0)
df.head()

```

```python
def crit(df_B, df_A):
    alpha = 0.05
    stat, pvalue_tcrit = stats.ttest_ind(df_B, df_A) #t-критерий Стьюдентa
    mannwhitneyu_stat, pvalue_mann = stats.mannwhitneyu(df_B, df_A) #Критерий Манна-Уитни
    wilcoxon_stat, pvalue_wilc = stats.wilcoxon(df_B, df_A) #Критерий Уилкоксона
    return pvalue_tcrit, pvalue_mann, pvalue_wilc

pval = crit(df["Before"], df["After"])
pval
#гипотеза об эффективности терапии не отвергается
```

### vs, mpg

```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(12, 6))
sns.boxplot(x='vs', y='mpg', data=df, ax=ax1)
ax1.set_xlabel('Тип двигателя', fontsize=12)
ax1.set_ylabel('Расход топлива', fontsize=12)
ax1.set_title('Расход топлива в зависимости от типа двигателя', fontsize=14)

sns.histplot(df, x='mpg', hue='vs', element='step', alpha=0.7, stat='density', ax=ax2)
ax2.set_xlabel('Расход топлива', fontsize=12)
ax2.set_title('Расход топлива в зависимости от типа двигателя', fontsize=14)
```

```python
stats.ks_2samp(df[df['vs'] == 0]['mpg'], df[df['vs'] == 1]['mpg'])
#гипотеза о нормальности распределений отвергается
```

```python
# H0: расход топлива не зависит от типа двигателя
# H1: расход топлива зависит от типа двигателя
```

```python
def crit(df_B, df_A):
    alpha = 0.05
    stat, pvalue_tcrit = stats.ttest_ind(df_B, df_A) #t-критерий Стьюдентa
    mannwhitneyu_stat, pvalue_mann = stats.mannwhitneyu(df_B, df_A) #Критерий Манна-Уитни
    return pvalue_tcrit, pvalue_mann

pval = crit(df[df['vs'] == 0]['mpg'], df[df['vs'] == 1]['mpg'])
#нулевая гипотеза отвергается -> расход топлива зависит от типа двигателя
```

### am, mpg

```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(12, 6))
sns.boxplot(x='am', y='mpg', data=df, ax=ax1)
ax1.set_xlabel('Тип коробки передач', fontsize=12)
ax1.set_ylabel('Расход топлива', fontsize=12)
ax1.set_title('Расход топлива в зависимости от типа коробки передач', fontsize=14)

sns.histplot(df, x='mpg', hue='am', element='step', alpha=0.7, stat='density', ax=ax2)
ax2.set_xlabel('Расход топлива', fontsize=12)
ax2.set_title('Расход топлива в зависимости от типа коробки передач', fontsize=14)
```

```python
stats.ks_2samp(df[df['am'] == 0]['mpg'], df[df['am'] == 1]['mpg'])
#гипотеза о нормальности распределений отвергается
```

```python
# H0: расход топлива не зависит от типа коробки передач
# H1: расход топлива зависит от типа коробки передач
```

```python
def crit(df_B, df_A):
    alpha = 0.05
    stat, pvalue_tcrit = stats.ttest_ind(df_B, df_A) #t-критерий Стьюдентa
    mannwhitneyu_stat, pvalue_mann = stats.mannwhitneyu(df_B, df_A) #Критерий Манна-Уитни
    return pvalue_tcrit, pvalue_mann

pval = crit(df[df['am'] == 0]['mpg'], df[df['am'] == 1]['mpg'])
pval
#нулевая гипотеза отвергается -> расход топлива зависит от типа коробки передач
```

### vs, am

```python
contingency_table = pd.crosstab(df['vs'], df['am'])
contingency_table
```

```python
contingency_table.plot(kind='bar')
```

```python
(contingency_table.T / contingency_table.sum(axis=1)).plot(kind='bar')
```

```python
# H0: тип передачи и тип двигателя независимы друг от друга
# H1: тип передачи и тип двигателя зависимы друг от друга

from scipy.stats import chi2_contingency
chi, pvalue, df_, expected = chi2_contingency(contingency_table.values, correction=False) #критерий Пирсона
#нулевая гипотеза не отвергается
```

```python
#через нормальное распределение
n = 32
c = 7
b = 6
z_obs = (b-c)/np.sqrt(b+c - (b-c)**2/n)
pval = 2*(1-stats.norm.cdf(z_obs))
pval
```